Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `???` or "YOUR ANSWER HERE".

---

# CSCI 3155 Assignment 2 : Recursion and Tail Recursion

This assignment asks you to write scala programs. 

**Restrictions** apply to each problem in terms of forbidden Scala features and API functions. Please read them carefully and ask for clarifications from the course staff over Piazza or during office hours if unsure.

Note: `???` indicates that there is a missing function or code fragment that needs to be filled in. In scala, 
it is also a macro that throws a `NotImplemented` exception. Make sure that you remove the `???` and replace it with the answer. 

Use the test cases provided to test them. You are also encouraged to write your own test cases to help debug your work. However, please delete any extra cells you may have created lest they break our autograder.

**Very Important:** Please run the cell that defines the functions `passed` and `testWithMessage` below whenever you restart the notebook.

### YOUR NAME HERE

```scala
// TEST HELPER

// FIRST RUN THIS CELL EVERY TIME YOU START THE NOTEBOOK
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage[T](v1: T, expected: T, testID: String) = { 
    println(s"Test $testID"); 
    println(s"\t Expected: $expected, your code returned: $v1")
    assert (v1 == expected, s"Test $testID FAILED.")
    println("\t Passed!")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 0 (10 points)

Write a _tail recursive_ implementation of the function `findAllOddIndices` that given a list `List`($e_0, \cdots, e_n$) returns a list of integers consisting of indices 
`List`($i_1, \ldots, i_k$)  corresponding  to all the _odd number_ elements $e_{i_1}, \ldots, e_{i_k}$ in the original list. Note that you must return a list in the increasing order of indices.

**Restrictions**

- Must be tail recursive.
- Must not use any list operations/API functions other than the following:
  - Cons (`::`)
  - `length`, `head`, `tail`, `reverse`.

<!-- grade_id: cell-c09fac6387b15cfc, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
final def findAllOddIndices(lst: List[Int], i:Int = 0, rev_lst:List[Int] = Nil): List[Int] = 
{
    if(lst == Nil)
    {
        return rev_lst.reverse;
    }
    else
    {
        val head_elt = lst.head;
        if(head_elt % 2 != 0)
        {
            findAllOddIndices(lst.tail, i+1, i::rev_lst);
        }
        else
        {
            findAllOddIndices(lst.tail, i+1, rev_lst);
        }
    }
} 
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36mfindAllOddIndices[39m
```

<!-- grade_id: cell-98cf9955eb5c985e, points: 10, grade: true, solution: false -->
```scala
def testMsg(lst: List[Int], lst_expected: List[Int], testName: String):Unit = {
    val lst_ret = findAllOddIndices(lst) // Call the student's code
    if (lst_ret == lst_expected){ // compare result against expected
        println(s"Test $testName passed") // let the good times roll.
    } else {
        println(s"Test $testName failed") //oops..Houston, we have a problem.
        println(s"Input: ${lst}, Your code: ${lst_ret}, Expected: ${lst_expected}")
        assert(false, "Assertion failed")
    }
}

val lst1 = List(0, 5, 2, 1, 3, 4, 7)
val lst1_expected = List(1, 3, 4, 6)
testMsg(lst1,  lst1_expected, "#1")

val lst2: List[Int] = Nil
val lst2_expected:List[Int] = Nil 
testMsg(lst2, lst2_expected, "#2")

val lst3: List[Int] = List(0, 2, 4, 6, 8, 10, 12, 14, 16)
val lst3_expected: List[Int] = Nil
testMsg(lst3, lst3_expected, "#3")

val lst4: List[Int] = List(1, 3, 5, 7, 9, 11, 13, 15)
val lst4_expected: List[Int] = (0 until 8).toList
testMsg(lst4, lst4_expected, "#4")

val lst5: List[Int] = List(0, 0, 0, 0, 0, -5) // Sorry :-)
val lst5_expected:List[Int] = List(5)
testMsg(lst5, lst5_expected, "#5")

passed(10)
```

Output:
```
Test #1 passed
Test #2 passed
Test #3 passed
Test #4 passed
Test #5 passed

*** Tests Passed (10 points) ***
defined [32mfunction[39m [36mtestMsg[39m
[36mlst1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m0[39m, [32m5[39m, [32m2[39m, [32m1[39m, [32m3[39m, [32m4[39m, [32m7[39m)
[36mlst1_expected[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m3[39m, [32m4[39m, [32m6[39m)
[36mlst2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mlst2_expected[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mlst3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m0[39m, [32m2[39m, [32m4[39m, [32m6[39m, [32m8[39m, [32m10[39m, [32m12[39m, [32m14[39m, [32m16[39m)
[36mlst3_expected[39m: [32mList[39m[[32mInt[39m] = [33mList[39m()
[36mlst4[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m3[39m, [32m5[39m, [32m7[39m, [32m9[39m, [32m11[39m, [32m13[39m, [32m15[39m)

```

## Problem 1 (15 points)

### A (5 points) Recursively compute binary representation

Write a function `convertToBinary` that inputs  a non-negative integer $n$, 
and returns its binary representation as a `String`. 
  - Use of var and loops forbidden. 
  - Inbuilt functions that directly perform conversion are also forbidden.
  - Try a recursive solution but no need to make it tail recursive for this part.

__Example:__ if input is 5, output should be "101".

Use the following scheme:
  - If number is 0 or 1, we can convert them into binary string "0" or "1", respectively.
  - If input number `n` is odd, recursively convert `(n-1)/2` into binary and append a `1` to the end of the string returned.
  - else  (input number `n` is even) convert `n/2` into binary and append `0` to end of the string returned.

For your convenience we have provided an "iterative solution" that uses a while loop. But __your__ solution must conform to the restrictions above.

```scala
def convertToBinaryLoop(n0: Int): String = {
    require(n0 >= 0)
    var n:Int = n0
    var str = ""
    if (n == 0){
        "0"
    } else {
        while (n > 0 ) {
            val bit = n % 2 
            str = bit.toString + str 
            n = n/2 // This is integer division
        }
        str
    }
}
```

Output:
```
defined [32mfunction[39m [36mconvertToBinaryLoop[39m
```

```scala
println(s"0 in binary = ${convertToBinaryLoop(0)}")
println(s"1 in binary = ${convertToBinaryLoop(1)}")
println(s"5 in binary = ${convertToBinaryLoop(5)}")
println(s"10 in binary = ${convertToBinaryLoop(10)}")
println(s"12 in binary = ${convertToBinaryLoop(12)}")
println(s"19 in binary = ${convertToBinaryLoop(19)}")
println(s"29 in binary = ${convertToBinaryLoop(29)}")
println(s"51 in binary = ${convertToBinaryLoop(51)}")
```

Output:
```
0 in binary = 0
1 in binary = 1
5 in binary = 101
10 in binary = 1010
12 in binary = 1100
19 in binary = 10011
29 in binary = 11101
51 in binary = 110011
```

<!-- grade_id: cell-b203e0766dddb81c, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def convertToBinary(num: Int):String = 
{
    if (num == 0)
    {
        return "0";
    }
    else if (num == 1)
    {
        return "1";
    }
    else
    {
        //formula for converting decmial to bianry: num % 2 + 10 * convertToBinary(num/2)
        if(num %2 != 0)
        {
           val remain = (num-1)/2;
            convertToBinary(remain) + "1";
        }
        else
        {
            val div = num/2;
             convertToBinary(div) + "0";
        }
    }
    
}
```

Output:
```
defined [32mfunction[39m [36mconvertToBinary[39m
```

<!-- grade_id: cell-e185be4282235b2a, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(convertToBinary(1) == "1", "(1)decimal must be 1 binary")
assert(convertToBinary(5) == "101", "5 must be 101 binary")
assert(convertToBinary(11) == "1011", "(11) must be 1011")
assert(convertToBinary(4100) == "1000000000100", "(4100) must be 1000000000100")
assert(convertToBinary(30108100) == "1110010110110100111000100", "(30108100) must be 1110010110110100111000100")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

### B ( 10 points) _Tail_ Recursive `convertToBinary`
Now, write a tail recursive version of the `convertToBinary` function, called `convertToBinaryTail`. 

Your function should be tail recursive and takes in a single integer argument and returns a String. 

**Hint** You can have more arguments to your tail recursive function but make sure to give them a default value. Alternatively, you can make an explicit helper function. 

**Restrictions** 

Use of var and loops forbidden. Inbuilt functions that directly perform conversion are also forbidden. Must be tail recursive.

<!-- grade_id: cell-1411df4e4ec22bdc, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
final def convertToBinaryTail(num:Int, i:String = ""): String =
{
    if (num == 0)
    {
        return i +"0";
    }
    else if (num == 1)
    {
        return (i + "1").reverse;
    }
    else
    {
        //formula for converting decmial to bianry: num % 2 + 10 * convertToBinary(num/2)
        if(num %2 != 0)
        {
            val remain = (num-1)/2;
            convertToBinaryTail(remain, i + "1");
        }
        else
        {
             val div = num/2;
             convertToBinaryTail(div, i + "0");
        }
    }
}

```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36mconvertToBinaryTail[39m
```

<!-- grade_id: cell-9d54c16fa3009152, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(convertToBinaryTail(1) == "1", "(1)decimal must be 1 binary")
assert(convertToBinaryTail(5) == "101", "5 must be 101 binary")
assert(convertToBinaryTail(11) == "1011", "(11) must be 1011")
assert(convertToBinaryTail(4100) == "1000000000100", "(4100) must be 1000000000100")
assert(convertToBinaryTail(30108100) == "1110010110110100111000100", "(30108100) must be 1110010110110100111000100")
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 2 (15 points)

### Part A (7 points)

Write a tail recursive function `sumUpEvenIndices` that given a list of integers `List`($e_0$, ..., $e_{n-1}$) returns the sum of all elements $e_0 + e_2 + \cdots $
at even numbered indices.

## Corner Cases
 - Empty list must return 0.
 - List with just one element must return that element. 
 
### Restrictions
  - Must be tail recursive.
  - Allowed list API functions/operations: `list.length`, `list.head`, `list.tail`. 
    - All other operations over lists disallowed.
  - Do not try to convert the list into an array or some other data structure.
  - Use of vars/loops is not allowed.
 

<!-- grade_id: cell-c6c77b82bcccd63f, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
final def sumUpEvenIndices(lst:List[Int], sum:Int = 0, idx:Int = 0 ): Int =
{
    if(idx == lst.length)
    {
        return sum;
    }
    else
    {
        if( idx % 2 == 0)
        {
            sumUpEvenIndices(lst, sum + lst(idx), idx+1);
        }
        else
        {
            sumUpEvenIndices(lst, sum, idx+1);
        }
    }
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36msumUpEvenIndices[39m
```

<!-- grade_id: cell-90640429de869647, points: 7, grade: true, solution: false -->
```scala
val t1 = sumUpEvenIndices(List(5, 10, 12, -3) )
assert( t1 == 17, s"Test1 failed: expected answer 17, got: $t1")

val t2 = sumUpEvenIndices(List() )
assert( t2 == 0, s"Test2 failed: expected answer 0, got: $t2")

val t3 = sumUpEvenIndices(List(-3) )
assert( t3 == -3, s"Test3 failed: expected answer -3, got: $t3")

val t4 = sumUpEvenIndices(List(0, 1, 0, 5, 0, 9, 0, 12, 0, -45, 0, -56) )
assert( t4 == 0, s"Test4 failed: expected answer 0, got: $t4")

val t5 = sumUpEvenIndices(List(11, 0) )
assert( t5 == 11, s"Test5 failed: expected answer 11, got: $t5")

passed(7)
```

Output:
```
*** Tests Passed (7 points) ***
[36mt1[39m: [32mInt[39m = [32m17[39m
[36mt2[39m: [32mInt[39m = [32m0[39m
[36mt3[39m: [32mInt[39m = [32m-3[39m
[36mt4[39m: [32mInt[39m = [32m0[39m
[36mt5[39m: [32mInt[39m = [32m11[39m
```

### Part B (8 points)

Write a tail recursive function `zipLists` that given two lists of integers `List(e1, ..., em)` and  `List(f1, ..., fn)`, computes the list of pairs `List( (e1, f1), ..., (ek, fk) )` where $k = \min(m, n)$.

The function takes in two arguments of type `List[Int]` and returns `List[(Int, Int)]`. You may have extra arguments but set them to default values or have an explicit helper function. 

## Corner Cases
 - If one of the lists is empty, it must return empty list.
 
### Restrictions
  - Must be tail recursive.
  - Allowed list API functions/operations: `list.length`, `list.head`, `list.tail`, `list.reverse` and `::` (cons or append an element to the front of a list).
    - All other operations over lists are __not__ allowed. In particular do not use `zip`, `:+` (append to end) or `++` (join/concatenate two lists). 
  - Do not try to convert the list into an array or some other data structure.
  - Use of vars/loops is not allowed.
 


<!-- grade_id: cell-a68d2b264a132205, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
final def zipLists(lst_1: List[Int], lst_2:List[Int], lst_f:List[(Int, Int)] = Nil, idx_1: Int = 0, idx_2: Int = 0): List[(Int, Int)] = 
{
    if (idx_1 == lst_1.length || idx_2 == lst_2.length)
    {
        return lst_f.reverse;
    }
    else
    {
        zipLists(lst_1, lst_2, lst_f = (lst_1(idx_1), lst_2(idx_2)) :: lst_f , idx_1+1, idx_2+1)
    }
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36mzipLists[39m
```

<!-- grade_id: cell-53f9fbd16982ad66, points: 8, grade: true, solution: false -->
```scala
val l1 = zipLists(Nil, List(5, 2, 1, 2, 3))
assert(l1 == Nil, s"Test 1 failed: expected : Nil, obtained: $l1")

val l2 = zipLists( List(5, 2, 1, 2, 3), Nil)
assert(l2  == Nil, s"Test 2 failed: expected : Nil, obtained: $l2")

val l3 = zipLists(List(10), List(5, 2, 1, 2, 3))
assert(l3 == List((10,5)) , s"Test 3 failed: expected : List((10,5)), obtained: $l3")

val l4 = zipLists(List(0, 1, 2, 3, 4), List(-5, 2, -1, 2, -3))
assert(l4 == List( (0, -5), (1, 2), (2, -1), (3, 2), (4, -3) ), s"Test 4 failed: expected : List( (0, -5), (1, 2), (2, -1), (3, 2), (4, -3) ) , obtained: $l4")

val l5 = zipLists(List(1, 2, 3), List(5, 2, 1, 2, 3))
assert(l5 == List((1, 5), (2, 2), (3, 1)), s"Test 5 failed: List((1, 5), (2, 2), (3, 1)) expected : , obtained: $l5")

passed(8)
```

Output:
```
*** Tests Passed (8 points) ***
[36ml1[39m: [32mList[39m[([32mInt[39m, [32mInt[39m)] = [33mList[39m()
[36ml2[39m: [32mList[39m[([32mInt[39m, [32mInt[39m)] = [33mList[39m()
[36ml3[39m: [32mList[39m[([32mInt[39m, [32mInt[39m)] = [33mList[39m(([32m10[39m, [32m5[39m))
[36ml4[39m: [32mList[39m[([32mInt[39m, [32mInt[39m)] = [33mList[39m(([32m0[39m, [32m-5[39m), ([32m1[39m, [32m2[39m), ([32m2[39m, [32m-1[39m), ([32m3[39m, [32m2[39m), ([32m4[39m, [32m-3[39m))
[36ml5[39m: [32mList[39m[([32mInt[39m, [32mInt[39m)] = [33mList[39m(([32m1[39m, [32m5[39m), ([32m2[39m, [32m2[39m), ([32m3[39m, [32m1[39m))
```

## Problem 3 (10 points)

###  Bisection method for solving equations 

Suppose we wanted to find the root for an equation :  $ f(x) = 0$, 

The bisection method starts with two bounds $\ell < u$ where $f(\ell), f(u)$ have opposite signs. I.e, either 
  - (A) $ f(\ell) < 0$ and $f(u) > 0$ or 
  - (B) $f(\ell) > 0$ and $f(u) < 0$.

At each iteration, it computes $m = (\ell + u) /2$ and checks the sign of $f(m)$. Depending on the sign of $f(m)$ it updates $\ell, u$ by setting : $\ell = m$ if $f(\ell), f(m)$ have the same signs or $u = m$ if $f(u), f(m)$ have the same signs. 

The method terminates under the following conditions: 

  - $  | f(m) | \leq \epsilon$ where $\epsilon$ is a given tolerance (typically $10^{-10}$). In this case, $m$ is the root and we declare success.
  - $ | u - \ell | \leq \epsilon $ but $f(m) > \epsilon$. In this case, we return $m$ but declare failure.
  
Here is a version of the bisection method implemented using a while loop for your convenience.

```scala
def bisectWhileLoop(f: Double => Double, 
                    l0: Double, u0: Double, 
                    epsilon: Double=1E-10): (Double, Boolean) = {
    var l = l0
    var u = u0
    require ( u - l >= epsilon)
    require ( (f(l) < 0.0 && f(u) > 0.0 ) || (f(l) > 0.0 && f(u) < 0.0))
    /* Note: we may have as well just said assert(f(l) * f(u) < 0.0) */
    var m = (l + u)/2.0
    while (math.abs(f (m)) > epsilon && (u - l) > epsilon) {
        if ( f(m) * f(l) > 0) { /* Does l, m have same sign */
            l = m
        } else {
            assert(f(m) * f(u) > 0) /* u, m must have the same sign */
            u = m
        }
        m = (l + u)/2
    }
    if (math.abs(f (m)) <= epsilon){
        (m, true) /* good enough to count as a solution to f(x) = 0 */
    } else {
        (m, false) /* Not good enough for the given tolerance. */
    }
}
```

Output:
```
defined [32mfunction[39m [36mbisectWhileLoop[39m
```

```scala
// Let's compute sqrt(2)
val f: Double => Double = x => (x*x - 2)
val l0 = 1.0
val u0 = 2.0
val (res, success) = bisectWhileLoop(f, l0, u0)
```

Output:
```
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd14$Helper$$Lambda$2184/0x0000000801391ed0@3b608e9f
[36ml0[39m: [32mDouble[39m = [32m1.0[39m
[36mu0[39m: [32mDouble[39m = [32m2.0[39m
[36mres[39m: [32mDouble[39m = [32m1.4142135623842478[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

```scala
// Let's compute cube-root(5)
val f: Double => Double = x => (x*x*x - 5)
val l0 = 1.0
val u0 = 2.0
val (res, success) = bisectWhileLoop(f, l0, u0)
```

Output:
```
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd15$Helper$$Lambda$2194/0x0000000801393648@c3110c0
[36ml0[39m: [32mDouble[39m = [32m1.0[39m
[36mu0[39m: [32mDouble[39m = [32m2.0[39m
[36mres[39m: [32mDouble[39m = [32m1.7099759466655087[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

```scala
//Solve the equation sin(x) - cos(x) = 0
val f: Double => Double = x => math.sin(x) - math.cos(x)
val l0 = 0.5
val u0 = 1.3
val (res, success) = bisectWhileLoop(f, l0, u0,1E-15)
println(s"Estimated value of pi: ${4*res}")
```

Output:
```
Estimated value of pi: 3.141592653589794
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd16$Helper$$Lambda$2200/0x00000008013c4e38@7b8eae66
[36ml0[39m: [32mDouble[39m = [32m0.5[39m
[36mu0[39m: [32mDouble[39m = [32m1.3[39m
[36mres[39m: [32mDouble[39m = [32m0.7853981633974485[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

Implement a tail recursive version of bisection called `bisectTail` that inputs
 - Function `f: Double => Double`
 - limits `l: Double`, `u: Double`
 - tolerance `epsilon:Double` that defaults to `1E-10`.
 
 It must return `(Double, Boolean)` where the first component is the approximate root and the second component of the tuple signals success/failure. 
 
   - Your function must be tail recursive.
   - Avoid the use of `var` and loops in your code. 
   - Do not use any API functions outside of those used in the while loop solution above such as `math.abs`.

<!-- grade_id: cell-4e0427b9f8dde074, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
final def bisectTail (f: Double => Double, l0: Double, u0: Double, epsilon: Double=1E-10, m: Double=(l0 + u0)/2.0): (Double, Boolean) = 
{
    require ((f(l0) < 0.0 && f(u0) > 0.0 ) || (f(l0) > 0.0 && f(u0) < 0.0))

    if(f(m) * f(l0) > 0)
    {
        bisectTail(f, m, u0, epsilon, (m+u0)/2.0);
    }
    else if (f(m).abs <= epsilon)
    {
        return (m, true);
    }
    else if ((u0 - l0).abs <= epsilon && f(m) > epsilon)
    {
        return (m, false);
    }
    else 
    {
        bisectTail(f, l0, m, epsilon, (m+l0)/2.0);
    }
    
        
    
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36mbisectTail[39m
```

<!-- grade_id: cell-5df0a19ea24b023c, points: 3, grade: true, solution: false -->
```scala
// Let's compute sqrt(2)
val f: Double => Double = x => (x*x - 2)
val l0 = 1.0
val u0 = 2.0
val (res, success) = bisectTail(f, l0, u0,1E-10)
assert (success && math.abs(res - 1.4142135) <= 1E-6)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd18$Helper$$Lambda$2207/0x00000008013c6e58@35127b45
[36ml0[39m: [32mDouble[39m = [32m1.0[39m
[36mu0[39m: [32mDouble[39m = [32m2.0[39m
[36mres[39m: [32mDouble[39m = [32m1.4142135623813377[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

<!-- grade_id: cell-df92d032bc5425df, points: 2, grade: true, solution: false -->
```scala
// Let's compute sqrt(5)
val f: Double => Double = x => (x*x - 5)
val l0 = 2.0
val u0 = 2.5
val (res, success) = bisectTail(f, l0, u0,1E-10)
assert(success && math.abs(res - 2.236067977) <= 1E-08)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd19$Helper$$Lambda$2215/0x00000008013c2b70@6f2ce45d
[36ml0[39m: [32mDouble[39m = [32m2.0[39m
[36mu0[39m: [32mDouble[39m = [32m2.5[39m
[36mres[39m: [32mDouble[39m = [32m2.2360679775010794[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

<!-- grade_id: cell-10a1539481c60d74, points: 2, grade: true, solution: false -->
```scala
//Solve the equation sin(x) - cos(x) = 0
val f: Double => Double = x => math.sin(x) - math.cos(x)
val l0 = 0.5
val u0 = 1.3
val (res, success) = bisectWhileLoop(f, l0, u0,1E-15)
println(s"Estimated value of pi: ${4*res}")
assert(success && math.abs(res - 0.78539816) <= 1E-07)
passed(2)
```

Output:
```
Estimated value of pi: 3.141592653589794

*** Tests Passed (2 points) ***
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd20$Helper$$Lambda$2223/0x00000008013c18a8@677f6448
[36ml0[39m: [32mDouble[39m = [32m0.5[39m
[36mu0[39m: [32mDouble[39m = [32m1.3[39m
[36mres[39m: [32mDouble[39m = [32m0.7853981633974485[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

<!-- grade_id: cell-88f1375c6182473b, points: 3, grade: true, solution: false -->
```scala
//Solve the equation sin(x) - cos(x) = 0.5
val f: Double => Double = x => math.sin(x) - math.cos(x) - 0.5
val l0 = 1.0
val u0 = 1.5
val (res, success) = bisectTail(f, l0, u0,1E-15)
assert(success && math.abs(math.sin(res) - math.cos(res) -0.5) <= 1E-09)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36mf[39m: [32mDouble[39m => [32mDouble[39m = ammonite.$sess.cmd21$Helper$$Lambda$2232/0x00000008013c91d0@61c50272
[36ml0[39m: [32mDouble[39m = [32m1.0[39m
[36mu0[39m: [32mDouble[39m = [32m1.5[39m
[36mres[39m: [32mDouble[39m = [32m1.1467652873041563[39m
[36msuccess[39m: [32mBoolean[39m = [32mtrue[39m
```

## That's All Folks!